Membrane potential modulates ERK activity and cell proliferation

Author response:

The following is the authors’ response to the original reviews.

Reviewer #1 (Public review):

Summary:

This is a contribution to the field of developmental bioelectricity. How do changes of resting potential at the cell membrane affect downstream processes? Zhou et al. reported in 2015 that phosphatidylserine and K-Ras cluster upon plasma membrane depolarization and that voltage-dependent ERK activation occurs when constitutive active K-RasG12V mutants are overexpressed. In this paper, the authors advance the knowledge of this phenomenon by showing that membrane depolarization up-regulates mitosis and that this process is dependent on voltage-dependent activation of ERK. ERK activity's voltage-dependence is derived from changes in the dynamics of phosphatidylserine in the plasma membrane and not by extracellular calcium dynamics.

Strengths:

Bioelectricity is an important field for areas of cell, developmental, and evolutionary biology, as well as for biomedicine. Confirmation of ERK as a transduction mechanism, and a characterization of the molecular details involved in control of cell proliferation, is interesting and impactful.

Weaknesses:

The functional cell division data need to be stronger. They show that increasing K+ increases proliferation and argue that since a MEK inhibitor (U0126) reduces proliferation in K+ treated cells, K+ induces cell division via ERK. But I don't see statistics to show that the rescue is significant, and I don't see a key U0126-only control. If the U0126 alone reduces proliferation, the combined effect wouldn't prove much.

We thank the reviewer for constructive feedback. We repeated the experiment including the U0126-only control (5K+U). We updated Fig.1, presenting the newly obtained data with statistical analysis.

Also, unless I'm missing something, it looks like every sample in their control has exactly the same number of mitotic cells. I understand that they are normalizing to this column, but shouldn't they be normalizing to the mean, with the independent values scattering around 1? It doesn't seem like it can be paired replicates since there are 6 replicates in the control and 4 replicates in one of the conditions?

We apologize for the unclear description. As the reviewer pointed out, the experiments were not paired replicates due to the limited number of conditions that can be conducted as a single experiment. To overcome this problem, we always included a control condition (i.e. 5K) based on which normalization was performed. This is the reason the data in 5K is always 1 and the sample size of 5K is the largest. Data include 100-900 mitotic cells within the imaging frame of 6 hrs. We re-wrote the figure legend (Fig1) and the main text, which hopefully clarified our experimental framework.

Reviewer #2 (Public review):

Sasaki et al. use a combination of live-cell biosensors and patch-clamp electrophysiology to investigate the effect of membrane potential on the ERK MAPK signaling pathway, and probe associated effects on proliferation. This is an effect that has long been proposed, but convincing demonstration has remained elusive, because it is difficult to perturb membrane potential without disturbing other aspects of cell physiology in complex ways. The time-resolved measurements here are a nice contribution to this question, and the perforated patch clamp experiments with an ERK biosensor are fantastic - they come closer to addressing the above difficulty of perturbing voltage than any prior work. It would have been difficult to obtain these observations with any other combination of tools.

However, there are still some concerns as detailed in specific comments below:

Specific comments:

(1) All the observations of ERK activation, by both high extracellular K+ and voltage clamp, could be explained by cell volume increase (more discussion in subsequent comments). There is a substantial literature on ERK activation by hypotonic cell swelling (e.g. https://doi.org/10.1042/bj3090013, https://doi.org/10.1002/j.1460-2075.1996.tb00938.x, among others). Here are some possible observations that could demonstrate that ERK activation by volume change is distinct from the effects reported here:

(i) Does hypotonic shock activate ERK in U2OS cells?

(ii) Can hypotonic shock activate ERK even after PS depletion, whereas extracellular K+ cannot?

(iii) Does high extracellular K+ change cell volume in U2OS cells, measured via an accurate method such as fluorescence exclusion microscopy?

(iv) It would be helpful to check the osmolality of all the extracellular solutions, even though they were nominally targeted to be iso-osmotic.

This is an important point. We conducted several experiments and provided explanations to rule out the possibility that ERK activation can be explained solely by cell volume change. We measured the osmolarity of all solutions used in this paper, which were 296-305 mOsm/L. This information was added to the Material and Methods section (line 387). Under our experimental conditions, ERK activation was not observed with hypotonic 70 % nor 50% osmolarity solution (Fig.S2).

It is therefore unlikely that the main cause of ERK activation upon high K+ perfusion is due to cell volume change. We would like to pursue this issue further when we obtain capacity to measure accurate cell volume change in the future.

(2) Some more details about the experimental design and the results are needed from Figure 1:

(i) For how long are the cells serum-starved? From the Methods section, it seems like the G1 release in different K+ concentration is done without serum, is this correct? Is the prior thymidine treatment also performed in the absence of serum?

Only the high K+ incubation phase was serum free. We added the following sentence in the main text (line 63) and an experimental diagram was added as Fig1A. “Cells were incubated in the presence of serum except for the phase with altered K+ concentration. “

(ii) There is a question of whether depolarization constitutes a physiologically relevant mechanism to regulate proliferation, and how depolarization interacts with other extracellular signals that might be present in an in vivo context.

This is a very important point. However, the significance of membrane depolarization for cell proliferation in vivo is beyond the scope of this study. This important question will be addressed in the future.

Does depolarization only promote proliferation after extended serum starvation (in what is presumably a stressed cell state)?

Cells were cultured in the presence of serum prior to the high K+ incubation phase as described above. We added a new figure (Fig1A).

What fraction of total cells are observed to be mitotic (without normalization), and how does this compare to the proliferation of these cells growing in serum-supplemented media? Can K+ concentration tune proliferation rate even in serum-supplemented media?

We included data recorded in serum-supplemented conditions (Fig.1), which showed a high mitotic rate. This is presumably due to the growth factors included in serum. There is no significant difference between 5K+FBS and 15K+FBS.

(3) In Figure 2, there are some possible concerns with the perfusion experiment:

(i) Is the buffer static in the period before perfusion with high K+, or is it perfused? This is not clear from the Methods. If it is static, how does the ERK activity change when perfused with 5 mM K+? In other words, how much of the response is due to flow/media exchange versus change in K+ concentration?

The buffer was static prior to high K perfusion. We confirmed that perfusion alone does not activate ERK (Fig.S2). We added the following sentence to the main text. “We also confirmed that the effect of perfusion was negligible, as ERK activation was not observed upon start of the 5K+ perfusion” (line 150).

(ii) Why do there appear to be population-average decreases in ERK activity in the period before perfusion with high K+ (especially in contrast to Fig. 3)? The imaging period does not seem frequent enough for photo bleaching to be significant.

Although we don’ t have a clear answer to this question, we speculate that several aspects of the experimental setup may have contributed to the difference. The cell lines and imaging systems used in Fig.2 and Fig.3 were different. The expression level may be different between U2OS cells and HEK 293 cells: transient expression in U2OS cells in contrast to stable expression in HEK 293 cells. This difference may lead to the different signal-to-noise ratio. The imaging system used in Fig.2 is an epi-illumination microscope excited with a 439/24 bandpass filter and detected with 483/32 (CFP) and 542/27 (YFP), while the imaging system used in Fig.3 is a confocal microscope excited with 458 nm laser and detected with 475-525 (DFP) and LP530 (YFP). These optical setups may also contribute to the different population-average properties before stimulation.

(4) Figure 3 contains important results on couplings between membrane potential and MAPK signaling. However, there are a few concerns:

(i) Does cell volume change upon voltage clamping? Previous authors have shown that depolarizing voltage clamp can cause cells to swell, at least in the whole-cell configuration: https://www.cell.com/biophysj/fulltext/S0006-3495(18)30441-7 . Could it be possible that the clamping protocol induces changes in ERK signaling due to changes in cell volume, and not by an independent mechanism?

We do not know whether cell volume is altered in the perforated-patch configuration. As discussed above, however, the effect of cell volume changes on ERK activity seemed to be negligible, because ERK activation was not observed with hypotonic 70 % nor 50% osmolarity solution (Fig.S2)

(ii) Does the -80 mV clamp begin at time 0 minutes? If so, one might expect a transient decrease in sensor FRET ratio, depending on the original resting potential of the cells. Typical estimates for resting potential in HEK293 cells range from -40 mV to -15 mV, which would reach the range that induces an ERK response by depolarizing clamp in Fig. 3B. What are the resting potentials of the cells before they are clamped to -80 mV, and why do we not see this downward transient?

We set the potential to -80mV immediately after the giga-seal formation and waited for at least 5 minutes to allow pore formation by gramicidin. We started imaging only after membrane potential was expected to have reached a steady state at -80 mV. We now included this sentence in the ‘Material and Methods’ section (line 398).

(5) The activation of ERK by perforated voltage clamp and by high extracellular K+ are each convincing, but it is unclear whether they need to act purely through the same mechanism - while additional extracellular K+ does depolarize the cell, it could also be affecting function of voltage-independent transporters and cell volume regulatory mechanisms on the timescales studied. To more strongly show this, the following should be done with the HEK cells where there is already voltage clamp data:

(i) Measure resting potential using the perforated patch in zero-current configuration in the high K+ medium. Ideally this should be done in the time window after high K+ addition where ERK activation is observed (10-20 minutes) to minimize the possibility of drift due to changes in transporter and channel activity due to post-translational regulation.

We measured membrane potential in the perforated patch configuration and confirmed that there is negligible potential drift within 20 minutes of perfusion with 145 K+ (only 1~5 mV change during perfusion).

(ii) Measure YFP/CFP ratio of the HEK cells in the high K+ medium (in contrast to the U2OS cells from Fig. 2 where there is no patch data).

YFP/CFP ratio data in HEK cells are shown in Fig.S1. As the signal-to-noise level is affected by the expression level of the probe, it is difficult to compare between cells with different expression levels. A higher YFP/CFP value with HEK cells compared to HeLa cells and A431 cells (Sup1) does not necessarily mean that HEK cells have higher ERK activity.

(iii) The assertion that high K+ is equivalent to changes in Vmem for ERK signaling would be supported if the YFP/CFP change from K+ addition is comparable to that induced by voltage clamp to the same potential. This would be particularly convincing if the experiment could be done with each of the 15 mM, 30 mM, and 145 mM conditions.

The experimental system using fluorescent biosensor cannot measure absolute ERK activity and can only measure the amount of change after a specific stimulus compared to the period before the stimulus. In electrophysiology experiments, the pre-stimulation membrane potential was clamped to -80 mV, whereas in the perfusion experiment, the membrane potential was variable in individual cells (-35 to -15 mV). It is therefore difficult to compare the results of electrophysiology experiments with those of the perfusion system. Unlike ion channels, it is currently not possible to plot absolute ERK activity with respect to the overall membrane potential. In the present study, we therefore discussed the change rather than the absolute value of ERK activity.

(6) Line 170: "ERK activity was reduced with a fast time course (within 1 minute) after repolarization to -80 mV." I don't see this in the data: in Fig. 3C, it looks like ERK remains elevated for > 10 min after the electrical stimulus has returned to -80 mV

Thank you for pointing out that our description was confusing. We changed the sentence to clarify the point we wanted to make. It now reads as follows. “ERK activity showed signs of reduction within 1 minute after repolarization to -80 mV.” (line 174)

Reviewer #3 (Public review):

Summary:

This paper demonstrates that membrane depolarization induces a small increase in cell entry into mitosis. Based on previous work from another lab, the authors propose that ERK activation might be involved. They show convincingly using a combination of assays that ERK is activated by membrane depolarization. They show this is Ca2+ independent and is a result of activation of the whole K-Ras/ERK cascade which results from changed dynamics of phosphatidylserine in the plasma membrane that activates K-Ras. Although the activation of the Ras/ERK pathway by membrane depolarization is not new, linking it to an increase in cell proliferation is novel.

Strengths

A major strength of the study is the use of different techniques - live imaging with ERK reporters, as well as Western blotting to demonstrate ERK activation as well as different methods for inducing membrane depolarization. They also use a number of different cell lines. Via Western blotting the authors are also able to show that the whole MAPK cascade is activated.

Weaknesses

A weakness of the study is the data in Figure 1 showing that membrane depolarization results in an increase of cells entering mitosis. There are very few cells entering mitosis in their sample in any condition. This should be done with many more cells to increase confidence in the results.

We apologize that that description was not clear. Due to the limited number of conditions that can be conducted as a single experiment, we always included control condition (i.e. 5K) and performed normalization by comparing with the control condition of the initial 1.5 hrs. Data were from 100-900 mitotic cell counts within 6hr of the imaging time window. We re-wrote the figure legend (Fig1) and the main text.

The study also lacks a mechanistic link between ERK activation by membrane depolarization and increased cell proliferation.

The present study focused on the link between membrane potential and the ERK activity; the mechanistic link between ERK activity and cell proliferation is beyond the scope of the present study. This important topic will be pursued further in subsequent studies.

The authors did achieve their aims with the caveat that the cell proliferation results could be strengthened. The results for the most part support the conclusions.

This work suggests that alterations in membrane potential may have more physiological functions than action potential in the neural system as it has an effect on intracellular signalling and potentially cell proliferation.

Reviewer #1 (Recommendations for the authors):

minor typo:

ERK activity has voltage-dependency with the physiological rang of membrane potential should be "range"

Corrected

Reviewer #2 (Recommendations for the authors):

Small points:

Line 82: rang -> range

Corrected

Line 102: ". they were stimulated" -> ". The cells were stimulated"

Corrected

Figs. 2C, 2D show exactly the same data points and the same information. Please cut one of these figures.

We deleted 2C and added the information in 2D and made new Fig.2C.

For all figs: Please indicate # of cells and # of independent dishes used in each experiment, and make clear whether individual data-points correspond to cells, dishes, or some other unit of measure.

We added the information in figure legends.

Reviewer #3 (Recommendations for the authors):

The authors should repeat the cell proliferation experiments with more cells to strengthen the data. They could also use alternative assays like phosphorylated histone H3 staining for cells in M phase, that might to easier to quantitate.

We repeated the experiment and Fig.1 was replaced with the new Fig.1

The authors should investigate how the upregulation of ERK is driving cells into mitosis. At what point in the cell cycle is activated ERK induced by membrane depolarization having the effect. Is it entry into mitosis or earlier in the cell cycle?

The cells were incubated with a high K+ solution 8-9 hr after G1 release, which is supposed to correspond to G2. These data suggest that mitotic activity is stimulated when ERK is activated at G2. However, we lack conclusive data at present to show the consequence of ERK activation during G2. We therefore cannot pinpoint the stage of cell cycle where depolarization-activated ERK exerts its effect.

The authors refer a lot to the work of Zhou et al 2015 throughout the paper. This is not necessary and is a bit distracting.

We deleted several sentence from the manuscript.

Changes from the earlier version to the revised version

Abstract

Summary

The plasma membrane potential has been linked to cell proliferation for more than 40 years. Here we experimentally showed that membrane depolarization upregulates cell mitosis, and that this process is dependent on voltage-dependent activation of ERK. ERK activity exhibits a membrane potential-dependency that is independent from the growth factor. This membrane potential dependence was observed even close to the resting membrane potential, indicating that small changes in resting membrane potential can alter cell proliferative activity. The voltage-dependent ERK activity is derived from changed dynamics of phosphatidylserine which is present in the plasma membrane and not by extracellular calcium entry. The data suggests that crucial biological processes such as cell proliferation are regulated by the physicochemical properties of the lipid. This study suggests that membrane potential may have diverse physiological functions beyond the action potential, which is well-established in the neural system.

Introduction

The plasma membrane potential is the voltage difference across the bilayer membrane of a cell that is established by the balance of intracellular and extracellular ion concentrations and their conductance. Intracellular and extracellular ion concentrations are maintained by passive and active ion transport through various ion channels and transporters located in the membrane. While individual cells have distinct resting plasma membrane potentials, the physiological function of membrane potential has been discussed mainly in excitable cells, i.e. muscle and nerve, and non-excitable cells have received less attention (1, 2).

However, accumulating evidence suggests that the physiological importance of membrane potential is not limited to the action potential in the excitable cells (3). Membrane potential has been suggested to be involved in pattern formation in early embryogenesis, regeneration and cancer (4). For example, membrane potential is a determinant factor for the left-right patterning in early embryogenesis (5). It is also known that resting membrane potential is related to proliferative capacity (6-10). These studies showed that depolarized cells have a higher level of mitotic activity. Moreover, Voorhess et al. reported that sustained depolarization induced DNA synthesis in chick spinal cord neurons (11). These results are consistent with the observation that oncogenetically transformed cells are generally more depolarized than normal parental cells (12). Although there have been many reports indicating that membrane potential regulates cell proliferation, the molecular mechanisms by which membrane potential regulates cell proliferation have not been investigated in parallel.

Earlier version: Cell proliferation is regulated by various growth factors and downstream intracellular signaling cascades such as mitogen-activated protein kinase (MAPK) cascades. Extracellular signalregulated kinase (ERK) is one of the key molecules that regulate cell proliferation (13). The ERK signaling pathway is often upregulated in human tumors. In contrast, when cells become confluent and proliferative activity becomes slow, ERK activity is downregulated (14).

Revised version: Cell proliferation is regulated by various growth factors and downstream intracellular signaling cascades such as mitogen-activated protein kinase (MAPK) cascades. Extracellular signal-regulated kinase (ERK) is one of the key molecules that regulate cell proliferation (13). The ERK signaling pathway is often upregulated in human tumors. In contrast, when cells become confluent and proliferative activity becomes slow, ERK activity is downregulated (14).

Earlier version: Zhou et al. reported that phosphatidylserine and K-Ras appear to undergo a change in their nanoscale organization upon plasma membrane depolarization (15). Zhou et al. also demonstrated that voltage-dependent ERK activation occurs when constitutive active K-RasG12V mutant are overexpressed. However, cell proliferative capacity was not the main focus of this paper. It still remains unclear whether voltage-dependent ERK can explain the membrane potential-dependent cell proliferative capacity.

Revised version: Phosphatidylserine and K-Ras have been reported to undergo a change in their nanoscale organization upon plasma membrane depolarization. Voltage-dependent ERK activation occurs when constitutive active K-RasG12V mutant are overexpressed (15). However, cell proliferative capacity was not the main focus of this paper. It still remains unclear whether voltage-dependent ERK can explain the membrane potential-dependent cell proliferative capacity.

Here, we show experimental evidence that depolarization does induce mitosis, which is associated with the activation of ERK.

Result

Mitotic activity is upregulated by membrane depolarization via MAPK cascade

Earlier version: Cell proliferation has been linked to membrane voltage for decades (9). We first investigated whether there is a causal relationship between membrane potential and mitotic activity in our experimental system. To focus on the mitotic phase of the cell cycle, U2OS cells were synchronized at the G1 phase by treating them with 2 mM thymidine for 20 hours, followed by 8-9 hours of normal incubation (G1 release). The cells were subsequently time-lapse imaged. Cytokinesis, which was easily recognized by brightfield images (Fig.1A), was counted for each frame. The number of cells in the mitotic phase was increased 12-13 hours after G1 release, which corresponds to 4 hours after the start of time-lapse imaging (Fig1B.). The mitotic activity, defined by the number of cells that have undergone cell division, was significantly upregulated in cells treated with a 15 mM high K+ solution (Fig.1B and C). This result indicates that membrane depolarization induce mitotic activity, which means there is a causal relationship between membrane potential and mitotic activity.

Revised version: Cell proliferation has been linked to membrane voltage for decades (9). We first investigated whether there is a causal relationship between membrane potential and mitotic activity in our experimental system. To focus on the mitotic phase of the cell cycle, U2OS cells were synchronized at the G1 phase by treating them with 2 mM thymidine for 20 hours, followed by 8-9 hours of normal incubation (G1 release). The cells were subsequently time-lapse imaged (Fig.1A). Cytokinesis, which was easily recognized by brightfield images (Fig.1B), was counted for each frame. Cells were incubated in the presence of serum except for the phase with altered K+ concentration (Fig.1A). The number of cells in the mitotic phase was increased 14-15 hours after G1 release, which corresponds to 6 hours after the start of time-lapse imaging (Fig.1C). Due to the limited number of conditions that can be conducted as a single experiment, the control condition (i.e. 5K) was always included for normalization between experiments. The mitotic activity, defined by the number of cells that have undergone cell division, was significantly upregulated in cells treated with a 15 mM high K+ solution (Fig.1C and D). This result indicates that membrane depolarization induce mitotic activity, which means there is a causal relationship between membrane potential and mitotic activity. Surprisingly, the mitotic activity in the 15K condition was comparable to that of in the FBS conditions (5K+FBS and 15K+FBS). The effects of depolarization and FBS were not additive, indicating that the proliferative activity was saturated in both condition.

Mitotic activity is upregulated by membrane depolarization.

Earlier version: (A) Representative image of mitotic cell. The images were 20 minute before and after cytokinesis. Scale bar is 10 μm. (B) The number of mitotic cells was accumulated for every 1.5 hours and normalized against those of first 1.5 hour with 5 K solution and plotted against time after starting the time-lapse imaging. Data are mean±SEM (C) Accumulated total mitotic cell number for 7.5 hours. p=0.044 (one sample t-test). U; U0126

Revised version: (A) Experimental diagram. (B) Representative image of mitotic cell. The images were 20 minute before and after cytokinesis. Scale bar is 10 μm. (C) Mitotic cell counts were accumulated every 1.5 hours, normalized by the number of mitotic cells in the first 1.5 hours with 5 K solution from the same experiment and plotted against time after the start of time-lapse imaging. U; U0126, FBS; 10% FBS. Data are mean±SEM (D) Mitotic cell counts were accumulated for 6 hours and normalized by the accumulated number of mitotic cells in 5 K solution from the same experiment. 5K v.s. 15K p=0.016 (one sample t-test), 5K v.s. 5K+FBS p=0.041 (one sample t-test), 5K v.s. 15K+FBS p=0.0094 (one sample t-test), 15K v.s. 15K+U p=0.0040 (paired t-test). Data are from 13 (for 5K), 13 (for 15K), 7 (for 5K+U), 11 (for 15K+U), 7 (for 5K+FBS) and 7 (for 15K+FBS) independent experiments respectively.100-900 mitotic cells were observed within 6 hours in each experiment.

Earlier version: We next investigated the potential involvement of ERK in the depolarization-induced mitotic activity. When cells were treated with MEK inhibitor (U0126), the effect of 15mM high K+ solution-induced mitotic activation was abolished, indicating that MEK activation play a pivotal role in depolarization induced proliferation (Fig.1B and C).

Revised version: We next investigated the potential involvement of ERK in the depolarization-induced mitotic activity. When cells were treated with MEK inhibitor (U0126), the effect of 15mM high K+ solution-induced mitotic activation was abolished, indicating that MEK activation play a pivotal role in depolarization induced proliferation (Fig.1C and D).

These result indicated that MEK-ERK activation was involved in the depolarization-induced mitotic activation. However, our results were not consistent with those of a previous report, which indicated that ERK is only activated when the extracellular potassium concentration is above 50 mM in cells expressing wild-type K-Ras (15). Therefore, it remains unclear whether the depolarization-induced mitotic activation could be explained by voltage-dependent activation of ERK. We therefore investigated whether ERK is activated by physiological range of membrane depolarization.

ERK activity has voltage-dependency with the physiological range of membrane potential

While Western blot is an excellent method for studying all-or-nothing phenomena, it is not well-suited for quantitatively capturing small changes or capturing phenomena with large cell to cell variability. Furthermore, it is not an appropriate method for studying detailed kinetics. To address these limitations, live cell imaging was used in this study. ERK activity was monitored by the Förster Resonance Energy Transfer (FRET) imaging using an ERK biosensor called EKAREV while the plasma membrane potentials were manipulated by changing the extracellular potassium concentration. EKAREV is the intramolecular FRET biosensor which changes its conformation depending on the ERK activity (16). Ypet and ECFP were used as the FRET pair, which we will refer to as YFP and CFP, respectively, for simplicity.

When U2OS cells expressing EKAREV were perfused with 145 mM high K+ solution, the YFP/CFP ratio increased over the next 10-20 min, indicating that ERK activity was upregulated (Fig.2A, B). This kinetics is consistent with previous report (15). We confirmed that our experimental system could capture depolarizing-induced ERK activation.

ERK is activated upon high K+ perfusion.

Earlier version: (A) YFP/CFP ratio images at the indicated time points. Serum starved U2OS cells expressing EKAREV were perfused with 145 mM high K+ solution at 10 minutes, followed by the addition of 10 nM EGF at 30 minutes. Scale bar is 10 mm. The color scale indicates YFP/CFP ratio. (B) ERK activity with high K+ perfusion. Serum-starved cells were treated with various concentrations of K+ solution at 10 min, followed by the addition of 10 nM EGF at 30 min. Normalized % ERK (response by high K / response by EGF; see in the text) activity was calculated as follows, the normalized YFP/CFP (see in the text) was divided by the normalized maximum YFP/CFP after subtracting the mean normalized YFP/CFP at 8-10 min. Data are mean±SEM. N=16, 15 and 23 for 145K+, 30K+ and 15K+, respectively. (C) The relationship between the K+ concentration and the % ERK activity. p=0.00013 (Kruskal-Wallis). (D) The relationship between the logarithm of K+ concentration and the % ERK activity. R is Pearson’s correlation coefficient. The gray area is 95% confidence interval. (E) Mean normalized YFP/CFP ratio perfused with Ca+ free 145 K+ high K solution. Starved cells were incubated with 5K+ (Ca+ free) solution, followed by perfusion with 145mM K+ (Ca+ free) solution at 10 min. 10nM EGF was added at 30 min. Data are mean±SEM. N=14 (F) The % ERK activity in the Ca(+) and Ca(-) solution. The data for Ca(+) were the same as in Fig.2C(145K+). p=0.82 (Wilcoxon rank sum exact test)

Revised version: (A) YFP/CFP ratio images at the indicated time points. Serum starved U2OS cells expressing EKAREV were perfused with 145 mM high K+ solution at 10 minutes, followed by the addition of 10 nM EGF at 30 minutes. Scale bar is 10 mm. The color scale indicates YFP/CFP ratio. (B) ERK activity with high K+ perfusion. Serum-starved cells were treated with various concentrations of K+ solution at 10 min, followed by the addition of 10 nM EGF at 30 min. Normalized % ERK (response by high K / response by EGF; see in the text) activity was calculated as follows, the normalized YFP/CFP (see in the text) was divided by the normalized maximum YFP/CFP after subtracting the mean normalized YFP/CFP at 8-10 min. Data are mean±SEM. Data are from 16, 15 and 23 cells from 3, 4 and 6 independent experiments, for 145K+, 30K+ and 15K+, respectively. (C) The relationship between the logarithm of K+ concentration and the % ERK activity. R is Pearson’s correlation coefficient. The gray area is 95% confidence interval. p=0.00013 (Kruskal-Wallis). N are the same as in (B). (D) Mean normalized YFP/CFP ratio perfused with Ca+ free 145 K+ high K solution. Starved cells were incubated with 5K+ (Ca+ free) solution, followed by perfusion with 145mM K+ (Ca+ free) solution at 10 min. 10nM EGF was added at 30 min. Data are mean±SEM. N=14 from 2 independent experiments. (F) The % ERK activity in the Ca(+) and Ca(-) solution. The data for Ca(+) were the same as in Fig.2C(145K+). p=0.82 (Wilcoxon rank sum exact test)

Earlier version: For quantitative analysis, two types of normalization were combined in order to facilitate the detection of small differences. One is to normalize the variation among dishes and the other is to normalize the variation among cells. Since the raw YFP/CFP value varied among dishes, presumably due to the difference in the expression level of the EKAREV probe and the background noise, the YFP/CFP value before application of the high K+ solution was averaged and used as a reference. The ratio of the raw YFP/CFP value to the reference value was defined as the normalized YFP/CFP value. The other is to normalize the variation among cells. they were stimulated with EGF (10 ng/mL) to strongly activate ERK activity, which was used as the maximum ERK activity for each cell. We defined the % ERK activity as the high K+-induced ERK activity (relative to the reference) divided by the maximum activity with EGF stimulation. These quantitative analyses were performed on imaged single cells (Fig.2B-F).

Revised version: For quantitative analysis, two types of normalization were combined in order to facilitate the detection of small differences. One is to normalize the variation among dishes and the other is to normalize the variation among cells. Since the raw YFP/CFP value varied among dishes, presumably due to the difference in the expression level of the EKAREV probe and the background noise, the YFP/CFP value before application of the high K+ solution was averaged and used as a reference. The ratio of the raw YFP/CFP value to the reference value was defined as the normalized YFP/CFP value. The other is to normalize the variation among cells. The cells were stimulated with EGF (10 ng/mL) to strongly activate ERK activity, which was used as the maximum ERK activity for each cell. We defined the % ERK activity as the high K+-induced ERK activity (relative to the reference) divided by the maximum activity with EGF stimulation. These quantitative analyses were performed on imaged single cells (Fig.2B-F).

Earlier version: ERK activity was upregulated by perfusion with 145 mM high K+ solution (Fig. 2A and B). We observed similar results in HEK293 cells, HeLa cells and A431 cells (fig.S1). The mean ERK activity with 145 mM high K+ perfusion was unexpectedly close to 100% (Fig.2B), indicating that depolarization induced by 145 mM K+ has a similar impact on cells as EGF stimulation. Since extracellular 145 mM high K+ solution is distant from a physiological condition, we next examined the 30 mM K+ perfusion, which did not result in ERK activation in the previous report (15). By using our imaging-based experimental system, we were able to detect ERK activation with 30 mM K+ perfusion. We could detect ERK activation even with 15 mM K+ perfusion where we observed mitotic upregulation, which means even a small depolarization within the physiological range could activate ERK, which in turn promotes cell division. Of note, U2OS cells expressed wild-type K-Ras but not an active mutant of K-Ras, which means voltage-dependent ERK activation occurs not only in tumor cells but also in normal cells. The average ERK activity with 30 mM K+ perfusion was about half of the maximum. The average ERK activity was further reduced when perfused with the 15 mM K+ solution (Fig.2B). These results indicated that ERK has voltage dependent activity not only at high membrane potential (depolarization) but also in a range close to the resting membrane potential. This is an observation that differs from previous studies (15).

Revised version: ERK activity was upregulated by perfusion with 145 mM high K+ solution (Fig. 2A and B). We observed similar results in HEK293 cells, HeLa cells and A431 cells (fig.S1). The mean ERK activity with 145 mM high K+ perfusion was unexpectedly close to 100% (Fig.2B), indicating that depolarization induced by 145 mM K+ has a similar impact on cells as EGF stimulation. Since extracellular 145 mM high K+ solution is distant from a physiological condition, we next examined the 30 mM K+ perfusion, which did not result in ERK activation in the previous report (15). By using our imaging-based experimental system, we were able to detect ERK activation with 30 mM K+ perfusion. We could detect ERK activation even with 15 mM K+ perfusion where we observed mitotic upregulation, which means even a small depolarization within the physiological range could activate ERK, which in turn promotes cell division. Of note, U2OS cells expressed wild-type K-Ras but not an active mutant of K-Ras, which means voltage-dependent ERK activation occurs not only in tumor cells but also in normal cells. The average ERK activity with 30 mM K+ perfusion was about half of the maximum. The average ERK activity was further reduced when perfused with the 15 mM K+ solution (Fig.2B). These results indicated that ERK has voltage dependent activity not only at high membrane potential (depolarization) but also in a range close to the resting membrane potential.

Another advantage of our imaging-based experimental system is that we could quantify the ERK activity in each cell. When individual cells were examined for the responses, the overall distribution was not Gaussian and a large proportion of cells perfused with 145 mM K+ solution showed % ERK activity of 90% or higher (Fig. 2C and D). On the other hand, % ERK activity of cells perfused with 30 mM high K+ solution was concentrated around 50 ∼ 60%, and the averaged response was lower. The variance of the response was larger in cells perfused with 15 mM high K+, with a significant proportion of cells showing no response (Fig. 2C and D).

The correlation coefficient between the % ERK activity and the logarithm of the extracellular K+ concentration was 0.54 (Fig.2D). Since the membrane potential is proportional to the logarithm of the extracellular K+ concentration based on the Nernst equation, these data indicated that ERK activity was voltage dependent not only at high membrane potentials but also in the physiological range of membrane potentials.

ERK is activated by various growth factors as well as by various stimuli such as neural excitation or mechanosensory input (17, 18), where the extracellular source of calcium entry is occasionally implicated. ERK activation in neurons was abolished when extracellular calcium was depleted (19). Calcium entry is also important in mechanosensory input. Mechanical stretch induces the opening of the stretch-activated Piezo1 channel and the calcium influx. We therefore tested whether extracellular calcium entry is involved in voltage-dependent ERK activation.

When cells were perfused with calcium-free145 mM K+ solution, ERK was activated to the same extent as the EGF stimulation, suggesting that calcium entry from the extracellular space is not necessary for voltage-dependent activation of ERK (Fig.2E and F).

Added in the revised version: There is substantial amount of literature on ERK activation by hypotonic cell swelling (20, 21). In our experimental conditions, there was no ERK activation with 70 % or 50% osmolarity (Fig.S2). We also confirmed that the effect of perfusion was negligible, as ERK activation was not observed upon start of the 5K+ perfusion (Fig.S2)

ERK activity depends on plasma membrane potential

ERK was only activated above -10mV with wild-type K-Ras in the previous study (15). In order to gain further insight into the relationship between voltage-dependent ERK activation and absolute membrane potential, we employed electrophysiology as a research tool.

ERK activity was monitored by EKAREV FRET imaging, while depolarization was achieved by the voltage clamping. Since U2OS cells were not amenable to the giga-seal formation required for patch clamping, we instead used HEK293 cells stably expressing the ERK biosensor EKAREV.

In a preliminary experiment, we performed the whole-cell configuration but failed to observe ERK activation even with EGF stimulation, suggesting that the intracellular component necessary for ERK activation was depleted by the whole-cell configuration. Thus, we next turned to the perforated patch-clamp technique using gramicidine (called perforated patch clamp), which keeps the intracellular component intact while allowing ions to pass through membrane pores.

HEK293 cells stably expressing EKAREV were depolarized from a holding potential of -80 mV. A control cell in the same optical view, whose membrane voltage was not manipulated, is shown in blue (Fig. 3A-F).

ERK activity depends on membrane voltage

Earlier version: (A to C) Mean normalized YFP/CFP ratio. Starved cells were depolarized from a holding potential of -80mV to depolarized membrane potentials (−40, -20 and 0 mV for A, B, and C, respectively) at 5 minutes, followed by repolarization to -80mV at 10 minute. 10nM EGF was perfused at 15 minutes. Data are mean±SD. N=4, 8 and 10 for -40, -20 and 0 mV, respectively. (D to F) Average normalized YFP/CFP ratio for the first 5 min (−80mV), maximum normalized YFP/CFP ratio at depolarized phase (−40, -20 and 0mV for A, B, and C, respectively), minimum normalized YFP/CFP ratio at repolarized phase (−80mV) and maximum normalized YFP/CFP ratio at EGF stimulated phase (EGF), respectively. Data are mean±SEM. The p-values were calculated using Student’s t test. (G) The relationship between membrane potential and % ERK activity. p=0.0042 (Kruskal-Wallis). R is Pearson’s correlation coefficient. The gray area is 95% confidence interval.

Revised version: (A to C) Mean normalized YFP/CFP ratio. Starved cells were depolarized from a holding potential of -80mV to depolarized membrane potentials (−40, −20 and 0 mV for A, B, and C, respectively) at 5 minutes, followed by repolarization to -80mV at 10 minute. 10nM EGF was perfused at 15 minutes. Data are mean±SD. N=4, 8 and 10 from 2, 5 and 6 independent experiments for -40, -20 and 0 mV, respectively. (D to F) Average normalized YFP/CFP ratio for the first 5 min (−80mV), maximum normalized YFP/CFP ratio at depolarized phase (−40, −20 and 0mV for A, B, and C, respectively), minimum normalized YFP/CFP ratio at repolarized phase (−80mV) and maximum normalized YFP/CFP ratio at EGF stimulated phase (EGF), respectively. Data are mean±SEM. The p-values were calculated using Student’s t test. (G) The relationship between membrane potential and % ERK activity. p=0.0042 (Kruskal-Wallis). R is Pearson’s correlation coefficient. The gray area is 95% confidence interval.

Depolarization to −40 mV did not upregulate ERK (Fig.3A and D). Depolarization from -80 mV to -20 mV significantly activated ERK (Fig. 3B and 3E). Depolarization to 0 mV caused the strongest activation of ERK, close to the maximum level (Fig. 3C and F).

In our experimental condition, HEK cells exhibited resting membrane potential ranging from -36 mV to -15 mV. This is in line with the lack of ERK activation at -40 mV.

Removed in the revised version: Again we showed voltage-dependent ERK activation within the membrane potential range which was not observed in the previous study (15).

Earlier version: Notably, ERK activity was reduced with a fast time course (within 1 minute) after repolarization to -80 mV, which was much faster kinetics than the previous report of 30 min (15). Faithful and swift change of the ERK activity coupled to the membrane potential is in contrast to its activation caused by EGF, which does not decline for more than 10 minutes (Fig. 3B, 3C,3E and 3F).

Revised version: Notably, ERK activity started showing reduction within 1 minute after repolarization to -80 mV, which was much faster kinetics than the previous report of 30 min (15). Faithful and swift change of the ERK activity coupled to the membrane potential is in contrast to its activation caused by EGF, which does not decline for more than 10 minutes (Fig. 3B, 3C,3E and 3F).

The correlation coefficient between % ERK activity and membrane potential was 0.54 (Fig.3G), again suggesting that ERK activity is voltage dependent even when close to the resting membrane potential. We also confirmed that depolarization (0 mV) alone was sufficient to activate ERK to a level comparable (88±35%) to the EGF stimulation (Fig. 3G).

The variance of % ERK activity was larger with the depolarization pulse at -20mV compared to 0 mV or -40mV (Fig.3G). We speculate that it is due to the range of resting membrane potentials we observed, between -36 mV to -15 mV: i.e. the polarity of the voltage changes from the resting potential to -20 mV is expected to be different in individual cells, which determined the ERK activation.

MAPK cascade was activated upon depolarization

We confirmed ERK activation upon high K+ perfusion by Western blot. 145 mM high K+ perfusion was used in this experiment, because 50 mM high K+ did not upregulate ERK activity with Western blot (15). U2OS cells were perfused with 145 mM high K+ solution and subsequently lysed with lysis buffer. ERK was indeed phosphorylated upon high K+ perfusion (Fig.4A and B).

Added in the revised version: We next examined the involvement of upstream molecules implicated in the ERK signaling cascade, namely the MAPK cascade with same condition (145 mM K+). MEK and c-Raf, located upstream of ERK in the canonical pathway, were also phosphorylated by the 145 mM high K+ perfusion (Fig. 4A and 4B). These results indicated that the MAPK cascade had a depolarization-activated pathway.

The molecular mechanism by which regulate depolarization-induced mitosis.

Earlier version: (A) Immunoblot analysis of cells perfused with 145K+ solution. (B) Quantification of relative mean band intensity of pERK, pMEK and pcRaf from independent experiments conducted as in (A). Data are mean±SD. (C) Ras activity was monitored by FLIM-FRET imaging. The donor (mTurquoise-GL) lifetime of U2OS cells expressing Raichu-Ras at the indicated time point. The experimental condition was the same as in Fig.2A. Data are mean±SEM. N= 17 (D) Schema of intermolecular FRET experiment. Phosphatidylserine clustering was monitored by FLIM-FRET imaging of U2OS cells co-expressing CFP-LactC2 and YFP-LactC2. (E) The fluorescence lifetime of CFP-LactC2. The fluorescence lifetime of CFP-LactC2 decreased when YFP-LactC2was co-expressed in the same cells. p=0.00080 (Welch Two Sample t-test) (F) The normalized fluorescence lifetime of CFP-LactC2 was measured at the indicated time point. The experimental condition was the same as in Fig.2A except the EGF stimulation. Data are mean±SEM. N=25 (G) CFP-LactC2 lifetime images at the indicated time points. The experimental condition is described above. Scale bar is 10 μm. The color scale indicates CFP-LactC2 lifetime. (H) Mean normalized YFP/CFP ratio from phosphatidylserine-depleted starved cells treated with 145 K+ solution at 10 minute, followed by the addition of 10nM EGF. Data are mean±SEM. N=25 (I) The % ERK activity with or without Fendiline treatment. The data for Fendiline(−) were the same as in Fig.2C(145K+). p=4.4e-9(Wilcoxon rank sum exact test).

Revised version: (A) Immunoblot analysis of cells perfused with 145K+ solution. (B) Quantification of relative mean band intensity of pERK, pMEK and pcRaf from independent experiments conducted as in (A). Data are mean±SD. (C) Ras activity was monitored by FLIM-FRET imaging. The donor (mTurquoise-GL) lifetime of U2OS cells expressing Raichu-Ras at the indicated time point. The experimental condition was the same as in Fig.2A. Data are mean±SEM. N= 17 from 4 independent experiments. (D) Schema of intermolecular FRET experiment. Phosphatidylserine clustering was monitored by FLIM-FRET imaging of U2OS cells co-expressing CFP-LactC2 and YFP-LactC2. (E) The fluorescence lifetime of CFP-LactC2. The fluorescence lifetime of CFP-LactC2 decreased when YFP-LactC2 was co-expressed in the same cells. p=0.00080 (Welch Two Sample t-test) N= 25 and 16 from 6 and 5 independent experiments, for CFP-LactC2 alone and co-expressing CFP-LactC2 and YFP-LactC2, respectively. (F) The normalized fluorescence lifetime of CFP-LactC2 was measured at the indicated time point. The experimental condition was the same as in Fig.2A except the EGF stimulation. Data are mean±SEM. N=25 from 6 independent experiments. (G) CFP-LactC2 lifetime images at the indicated time points. The experimental condition is described above. Scale bar is 10 μm. The color scale indicates CFP-LactC2 lifetime. (H) Mean normalized YFP/CFP ratio from phosphatidylserine-depleted starved cells treated with 145 K+ solution at 10 minute, followed by the addition of 10nM EGF. Data are mean±SEM. N=27 from 8 independent experiments. (I) The % ERK activity with or without Fendiline treatment. The data for Fendiline(-) were the same as in Fig.2C(145K+). p=4.4e-9(Wilcoxon rank sum exact test).

Removed in the revised version: We next examined the involvement of upstream molecules implicated in the ERK signaling cascade, namely the MAPK cascade with same condition (145 mM K+). MEK and c-Raf, located upstream of ERK in the canonical pathway, were also phosphorylated by the 145 mM high K+ perfusion (Fig. 4A and 4B). These results indicated that the MAPK cascade had a depolarization-activated pathway.

Earlier version: Our real-time imaging data revealed that depolarization-induced ERK activation has considerably slower kinetics than voltage-dependent ion channels, which typically activate within a millisecond (20). We next examined the kinetics of the upstream pathway of MAPKs with imaging-based methods so that detail kinetics can be monitored.

Revised version: Our real-time imaging data revealed that depolarization-induced ERK activation has considerably slower kinetics than voltage-dependent ion channels, which typically activate within a millisecond (22). We next examined the kinetics of the upstream pathway of MAPKs with imaging-based methods so that detail kinetics can be monitored.

Ras was activated upon depolarization with fast kinetics

Earlier version: We examined Ras activity, which is upstream of c-Raf. Although nonoclusters of Ras were known to cause activation of downstream MAPK cascade, there were no direct data to date on whether Ras is activated by depolarization. Ras activity was monitored by imaging with a FRET-based biosensor called Raichu-Ras (21). Raichu-Ras is an intramolecular FRET biosensor which changes its conformation and triggers an increase in FRET depending on the Ras activity. The lifetime of the donor becomes shorter when FRET occurs. We aimed to monitor Ras activity in an area restricted to the vicinity of the plasma membrane. Ras activity was monitored by fluorescence lifetime imaging microscopy combined with fluorescence resonance energy transfer (FLIM-FRET) (21). U2OS cells expressing Raichu-Ras were perfused with 145mM high K+ solution followed by perfusion with EGF solution (10ng/mL).

Revised version: We examined Ras activity, which is upstream of c-Raf. Although nonoclusters of Ras were known to cause activation of downstream MAPK cascade, there were no direct data to date on whether Ras is activated by depolarization. Ras activity was monitored by imaging with a FRET-based biosensor called Raichu-Ras (23). Raichu-Ras is an intramolecular FRET biosensor which changes its conformation and triggers an increase in FRET depending on the Ras activity. The lifetime of the donor becomes shorter when FRET occurs. We aimed to monitor Ras activity in an area restricted to the vicinity of the plasma membrane. Ras activity was monitored by fluorescence lifetime imaging microscopy combined with fluorescence resonance energy transfer (FLIM-FRET) (23). U2OS cells expressing Raichu-Ras were perfused with 145mM high K+ solution followed by perfusion with EGF solution (10ng/mL).

Earlier version: The mean fluorescence lifetime of the donor fluorophore of Raichu-Ras on the plasma membrane decreased within 2 minutes after 145mM K+ perfusion, indicating the Ras activation (Fig. 4C). The activation kinetics were similar between high K+ stimulation and EGF stimulation, which was much faster than depolarization-induced ERK activation. The time course of activation was also consistent with the previous report (15, 22). This data suggested that Ras was activated by membrane depolarization with faster kinetics than that of ERK.

Revised version: The mean fluorescence lifetime of the donor fluorophore of Raichu-Ras on the plasma membrane decreased within 2 minutes after 145mM K+ perfusion, indicating the Ras activation (Fig. 4C). The activation kinetics were similar between high K+ stimulation and EGF stimulation, which was much faster than depolarization-induced ERK activation. The time course of activation was also consistent with the previous report (15, 24). This data suggested that Ras was activated by membrane depolarization with faster kinetics than that of ERK.

The phosphatidylserine dynamics was involved in the voltage-dependent ERK activity

Earlier version: We further investigated upstream of Ras. Change of membrane potential is expected to be complete within milliseconds, and kinetics upstream of Ras are expected to be at least as fast as that of Ras. Zhou et al. reported that plasma membrane depolarization induces nanoscale reorganization of phosphatidylserine (15). K-Ras, which is targeted to the plasma membrane by electrostatic interactions with phosphatidylserine, in turn undergoes nanoclustering and amplifies MAPK signaling in fibroblasts, excitable neuroblastoma cells, and Drosophila neurons (15). Based on these findings, we examined the involvement of phosphatidylserine nanoclustering in the voltage-dependent activation of ERK in U2OS cells.

Revised version: We further investigated upstream of Ras. Change of membrane potential is expected to be complete within milliseconds, and kinetics upstream of Ras are expected to be at least as fast as that of Ras. Plasma membrane depolarization has been reported to induce nanoscale reorganization of phosphatidylserine (15). K-Ras, which is targeted to the plasma membrane by electrostatic interactions with phosphatidylserine, in turn undergoes nanoclustering and amplifies MAPK signaling in fibroblasts, excitable neuroblastoma cells, and Drosophila neurons (15). Based on these findings, we examined the involvement of phosphatidylserine nanoclustering in the voltage-dependent activation of ERK in U2OS cells.

Earlier version: Phosphatidylserine dynamics were monitored by intermolecular FRET using the phosphatidylserine-specific probe LactC2 (Fig.4D). The C2 domain of Lactadherin (LactC2), which binds phosphatidylserine, allows the visualization of phosphatidylserine when conjugated with CFP or YFP (23). The fluorescence lifetime of CFP-LactC2 was 2.375±0.014 ns, which decreased to 2.345±0.013 ns when YFP-LactC2 was co-expressed in the same cells (Fig. 4E), indicating that phosphatidylserine nanoclustering could be monitored by FLIM-FRET. The fluorescence lifetime of CFP-LactC2 in the cells co-expressing the FRET acceptor YFP-LactC2 decreased within 2 minutes upon 145mM high K+ perfusion, indicating the depolarization-induced nanoclustering of phosphatidylserine (Fig. FC, 4G). The kinetics of depolarization-induced nanoclustering of phosphatidylserine is similar to that of Ras activation, which is a reasonable result if phosphatidylserine and Ras interact.

Revised version: Phosphatidylserine dynamics were monitored by intermolecular FRET using the phosphatidylserine-specific probe LactC2 (Fig.4D). The C2 domain of Lactadherin (LactC2), which binds phosphatidylserine, allows the visualization of phosphatidylserine when conjugated with CFP or YFP (25). The fluorescence lifetime of CFP-LactC2 was 2.375±0.014 ns, which decreased to 2.345±0.013 ns when YFP-LactC2 was co-expressed in the same cells (Fig. 4E), indicating that phosphatidylserine nanoclustering could be monitored by FLIM-FRET. The fluorescence lifetime of CFP-LactC2 in the cells co-expressing the FRET acceptor YFP-LactC2 decreased within 2 minutes upon 145mM high K+ perfusion, indicating the depolarization-induced nanoclustering of phosphatidylserine (Fig. FC, 4G). The kinetics of depolarization-induced nanoclustering of phosphatidylserine is similar to that of Ras activation, which is a reasonable result if phosphatidylserine and Ras interact.

Earlier version: To confirm that phosphatidylserine dynamics were involved in voltage-dependent ERK activation, phosphatidylserine was depleted by treating cells with fendiline for 48 hours (24). When U2OS cells were treated with fendiline for 36 to 48 hours, high K+-induced ERK activation was diminished (Fig. 4H and 4I). These results indicate that membrane phosphatidylserine dynamics play a critical role in voltage-dependent ERK activation in U2OS.

Revised version: To confirm that phosphatidylserine dynamics were involved in voltage-dependent ERK activation, phosphatidylserine was depleted by treating cells with fendiline for 48 hours (26). When U2OS cells were treated with fendiline for 36 to 48 hours, high K+-induced ERK activation was diminished (Fig. 4H and 4I). These results indicate that membrane phosphatidylserine dynamics play a critical role in voltage-dependent ERK activation in U2OS.

Discussion

The relationship between cell proliferation and membrane potential has been discussed since the 1970s. However, the lack of studies that addressed the molecular mechanisms has left the full picture of the relationship between membrane potential and cell proliferative potential unclear. With regards to the molecular mechanisms involved, a study published in 2015 demonstrated a correlation between MAP kinases activity and membrane potential. However, the voltage-dependence of ERK activity has not been demonstrated within the physiological membrane potential range. Therefore, it remains unclear whether voltage-dependent MAP kinase activation can explain the link between membrane potential and cell proliferation. In this study we were able to experimentally demonstrate a causal relationship between membrane potential and mitotic activity (cytokinesis) near the physiological membrane potential, which was associated with depolarization-induced ERK activation via phosphatidylserine dynamics. We were able to connect and clarify the molecular mechanism of the phenomenon that has long been suggested as a correlation between membrane potential and cell proliferative capacity.

Earlier version: We showed membrane depolarization induces mitotic activity through the activation of ERK (Fig.1), which indicates that membrane potential and cell proliferation are not merely correlated but membrane potential itself regulates mitotic activity. Some reports indicated that ERK plays an important role in the S phase entry (25), while other reports proposed that ERK plays a role in mitosis (26-28). Further studies will be needed to elucidate which step of the cell cycle is regulated by membrane potential.

Revised version: We showed membrane depolarization induces mitotic activity through the activation of ERK (Fig.1), which indicates that membrane potential and cell proliferation are not merely correlated but membrane potential itself regulates mitotic activity. Some reports indicated that ERK plays an important role in the S phase entry (27), while other reports proposed that ERK plays a role in mitosis (28-30). Further studies will be needed to elucidate which step of the cell cycle is regulated by membrane potential.

We showed direct evidence that membrane potential itself regulates ERK activity by controlling membrane potential using voltage-clamp to exclude any possibility that factors other than membrane potential regulate ERK activity. Another advantage of electrophysiology is the ability to study the effect of repolarization. With conventional techniques of perfusion it is difficult to change solution instantaneously, which hinders investigating the effect of repolarization.

Earlier version: The characteristics of the ERK response to voltage compared to EGF stimulation was the linearity of the stimulus and the response (Fig2 and 3). It is notable that ERK activity was downregulated with membrane repolarization, while ERK remains active for at least 5 minutes after the growth factor binds to its receptor (29). This implies that membrane potential can act as a fine tuner of ERK, in contrast to the growth factor stimulation acting as an on/off switch.

Revised version: The characteristics of the ERK response to voltage compared to EGF stimulation was the linearity of the stimulus and the response (Fig2 and 3). It is notable that ERK activity was downregulated with membrane repolarization, while ERK remains active for at least 5 minutes after the growth factor binds to its receptor (31). This implies that membrane potential can act as a fine tuner of ERK, in contrast to the growth factor stimulation acting as an on/off switch.

The imaging analysis using FRET-based probes allowed us to study kinetics for each step of signal transduction e.g. nanoclustering of phosphatidylserine and subsequent activation of Ras was completed within 2 minutes (Fig.4C, 4F), whereas activation of ERK took more than 5 minutes (Fig.2B). This result is reasonable since several molecules have to be activated before ERK activation.

Another advantage of imaging-based analysis is that it allowed us to quantitatively measure ERK activity in individual cells. The variability of the response to depolarization among analyzed cells was surprisingly large. Of note, this variance is averaged and not observable in Western blot analysis, which is routinely used to study intracellular signaling pathways, including ERK. Surprisingly, 0 mV depolarization activated ERK to the same extent as EGF stimulation.

Furthermore, coupled with electrophysiology, we could control and monitor the activity with a time resolution that cannot be achieved by perfusion or biochemical analysis. It revealed that ERK began to activate within 1 minute upon depolarization and sustained activity increase for at least 5 minutes. The kinetics of the deactivation upon repolarization was similar to that of activation.

We also demonstrated that ERK exhibits voltage-dependent behavior, ranging from the resting membrane potential to a high membrane potential (depolarization). This is reasonable when considering that voltage-dependent ERK activation is derived from voltage-dependent phosphatidylserine dynamics on the plasma membrane, which is supposed to be a physiochemical phenomenon. The plasma membrane potential is a novel member of the ERK regulator.

Earlier version: Zhou et al. reported that phosphatidylserine and PIP2 change their nanoscale organization upon plasma membrane depolarization (15). K-Ras also changes its nanoscale organization based on electrostatic interaction between its basic amino acid residues and anionic phosphatidylserine (15, 30). Such electrostatic interaction may occur between a cluster of basic amino acids in another protein and an anionic phospholipid, i.e. phosphatidylserine or PIP2. For example, it is known that a cluster of basic amino acids in EGFR interacts with PIP2 and that EGFR is activated when clustered (31, 32). This suggests that EGFR activity may also be regulated by membrane potential. There are many other proteins which interact with PIP2 and possess a cluster of basic amino acids, suggesting that membrane potential may regulate the activity of many more proteins than previously thought, and thereby regulate many physiological phenomena.

Revised version: Zhou et al. reported that phosphatidylserine and PIP2 change their nanoscale organization upon plasma membrane depolarization (15). K-Ras also changes its nanoscale organization based on electrostatic interaction between its basic amino acid residues and anionic phosphatidylserine (15, 32). Such electrostatic interaction may occur between a cluster of basic amino acids in another protein and an anionic phospholipid, i.e. phosphatidylserine or PIP2. For example, it is known that a cluster of basic amino acids in EGFR interacts with PIP2 and that EGFR is activated when clustered (33, 34). This suggests that EGFR activity may also be regulated by membrane potential. There are many other proteins which interact with PIP2 and possess a cluster of basic amino acids, suggesting that membrane potential may regulate the activity of many more proteins than previously thought, and thereby regulate many physiological phenomena.

Earlier version: We showed that membrane potential modulates phosphatidylserine dynamics (Fig.4F amd G). The relationship between the electrostatic potential and the organization of the cell membrane was previously discussed (33). The transmembrane potential affects the lateral sorting of membrane components as well as the positioning of polar lipid headgroups (33). In addition, lipid viscosity, which is influenced by lipid and protein composition and their organization in the membrane, is altered by the membrane potential (34-36). However, the exact mechanism by which the transmembrane potential can modulate the organization of lipids and proteins in the membrane remains unknown.

Revised version: We showed that membrane potential modulates phosphatidylserine dynamics (Fig.4F amd G). The relationship between the electrostatic potential and the organization of the cell membrane was previously discussed (35). The transmembrane potential affects the lateral sorting of membrane components as well as the positioning of polar lipid headgroups (35). In addition, lipid viscosity, which is influenced by lipid and protein composition and their organization in the membrane, is altered by the membrane potential (36-38). However, the exact mechanism by which the transmembrane potential can modulate the organization of lipids and proteins in the membrane remains unknown.

Earlier version: Overall, our results link three elements: i.e, the plasma membrane voltage, ERK activity and the cell proliferation, and propose a new signaling cascade regulating cell division. This result provides rationale that ion channels and ion transporters, which regulate the plasma membrane potential, are potential targets for tumor treatment. Moreover, given that ERK is involved also in cell differentiation, migration, senescence and apoptosis (37), membrane potential may also regulate them and prove to be potential targets for future drug developments.

Revised version: Overall, our results link three elements: i.e, the plasma membrane voltage, ERK activity and the cell proliferation, and propose a new signaling cascade regulating cell division. This result provides rationale that ion channels and ion transporters, which regulate the plasma membrane potential, are potential targets for tumor treatment. Moreover, given that ERK is involved also in cell differentiation, migration, senescence and apoptosis (39), membrane potential may also regulate them and prove to be potential targets for future drug developments.

Earlier version: What is the biological meaning of membrane potential regulating cell proliferation ? The stochastic ERK activation reported by Aoki et al. may result from the fluctuation of the membrane potential (38). Furthermore, cell-to-cell propagation of the membrane potential through gap junctions may result in spreading of the ERK activity and synchronized cell proliferation (38). This type of synchronized proliferation could be beneficial in tissue such as epithelium. Our finding that the plasma membrane potential is an analog modulator of ERK activity is consistent with this hypothesis, which would allow nuanced regulation of cell proliferation.

Revised version: What is the biological meaning of membrane potential regulating cell proliferation ? The stochastic ERK activation reported by Aoki et al. may result from the fluctuation of the membrane potential (40). Furthermore, cell-to-cell propagation of the membrane potential through gap junctions may result in spreading of the ERK activity and synchronized cell proliferation (40). This type of synchronized proliferation could be beneficial in tissue such as epithelium. Our finding that the plasma membrane potential is an analog modulator of ERK activity is consistent with this hypothesis, which would allow nuanced regulation of cell proliferation.

It is now evident that the physiological significance of membrane potential extends beyond its role in ion channels. Given that membrane potential is a physical phenomenon that affects all cellular components, including proteins and lipids, new biological phenomena regulated by membrane potential may exist.

Added in the revised version: Material and methods

Added in the revised version: Plasmid and reagents

Added in the revised version: The expression vectors for EKAREV and Raichu-Ras plasmid were kindly provided by Dr. Michiyuki Matsuda (Kyoto University). mRFP-Lact-C2 was a gift from Sergio Grinstein (Addgene plasmid # 74061 ; http://n2t.net/addgene:74061 ; RRID:Addgene_74061)(25). YFP-LactC2 and CFP-LactC2 plasmids were generated from mRFP-LactC2 plasmid by exchanging mRFP to YPet and SECFP, respectively. Recombinant human EGF was purchased from ThermoFisher Scientific. Fendiline was purchased from CAYMAN Chemical. Gramicidin was purchased from Sigma-Aldrich.

Added in the revised version: Cells

Added in the revised version: U2OS cells were purchased from ECACC (European Collection of Authenticated Cell Culture). U2OS cells were maintained in McCoy’s 5A medium supplemented with 10 % fetal bovine serum. HEK293 cells stably expressing EKAREV were generated by PiggyBac transposon system transfection with EKAREV/pPBbsr2 and mPBase/pCMV vectors and were maintained in DMEM medium supplemented with 10 % fetal bovine serum. Cells were cultured in a 5% CO2-humidified environment at 37°C.

Added in the revised version: Time-lapse imaging with high K+ perfusion U2OS cells were plated on the glass-based dish (Matsunami) and transfected with the EKAREV plasmid using TransIT-LT1 reagent (Mirus) according to the manufacture’s instructions. Cells were serum starved for 3 to 16 h before starting the time-lapse FRET imaging. Cells were imaged using an inverted microscope (Ti2-E; Nikon) equipped with a 60× objective (NA=1.4; Nikon), a CMOS camera (Zyla4.2; Andor), and an LED illumination system (xCite Xylis; Excelitas technologies). A W-VIEW GEMINI image splitting system (Hamamatsu Photonics) was used to acquire CFP and YFP images simultaneously. The following filters and dichroic mirrors were used: 439/24 for excitation, 483/32 (CFP) and 542/27 (YFP) for emission, and FF509-fDi01 for the dichroic mirror. FLIM-FRET images were acquired with a Stellaris (Leica) confocal microscope excited at 448 nm. The fluorescence lifetime of the donor fluorophore was measured at 460-500 nm with a pinhole size of 1 A.U. so that only signal from the plasma membrane was detected. The perfusion solution was changed using the valve-controlled gravity perfusion system (VC3-4PG) controlled by Arduino. The solutions were as follows (in mM), 5K ; 140 NaCl, 5 KCl, 1 CaCl2, 1 MgCl2. 15K ; 130 NaCl, 15 KCl, 1 CaCl2, 1 MgCl2. 30K ; 115 NaCl, 30 KCl, 1 CaCl2, 1 MgCl2. 145K ; 145 KCl, 1 CaCl2, 1 MgCl2. 5K (Ca free) ; 140 NaCl, 5 KCl, 1 EGTA, 1 MgCl2. 145K (Ca free) ; 145 KCl, 1 EGTA 1 MgCl2. All solutions contained 20 mM HEPES pH=7.4. Osmolarity was adjusted to 296 - 305 mOsm/L by glucose. EGF solution was prepared with low K+ solution (5 mM K+).

Added in the revised version: Patch clamping

Added in the revised version: HEK293 cells stably expressing EKAREV were plated on the glass-based dish (Matsunami) and serum starved for 16 hours before starting the experiment. The gramicidin-perforated patch clamp experiment was performed using a Multiclamp700B amplifier and Digidata 1440A and pClamp software (Axon). 60 μg/mL gramicidin in patch solution (145 mM KCl, 5 mM EGTA, 10 mM HEPES pH7.4) was sonicated for 5 minutes and back-filled to the patch pipette. 5K solution was used as external patch solution. FRET images were acquired with a confocal microscope LSM510Meta (Ziess) excited at 458 nm. The filters used for CFP and YFP were 475-525 and LP530, respectively. We set the potential to -80mV immediately after the giga-seal formation and waited for at least 5 minutes to allow pore formation by gramicidin. We started imaging only after membrane potential was expected to have reached a steady state at -80 mV.

Added in the revised version: Imaging Analysis

Added in the revised version: The dead cells and unhealthy cells that did not respond to EGF stimulation were excluded from the analysis. The average intensity of the multiple background ROIs was used as the background and subtracted from each image. The fluorescence intensities of the YFP channel and CFP channel were averaged for each cell, and the results were exported as a CSV file. The YFP/CFP value was calculated using Excel and R Studio for the patch-clamp and high-K+ perfusion experiments, respectively. The % ERK activity was calculated by dividing the maximum normalized YFP/CFP ratio after high K+ perfusion or depolarization by the maximum normalized YFP/CFP ratio after EGR perfusion for each cell. For the FLIM-FRET experiment, the average fluorescence lifetime was calculated for each cell. The fluorescence lifetime of the first 10 minutes was averaged and used as a reference. The ratio of fluorescence lifetime to the reference value was defined as normalized fluorescence lifetime. All image analysis was performed using ImageJ software.

Added in the revised version: Western blot

Added in the revised version: The following antibodies were used for Western blotting: total ERK (SantaCruz, SC514302); phospho-ERK (CST, 4370); total MEK (CST, 4394); phospho-MEK (CST, 2338); total c-Raf (CST, 9422); phospho-c-Raf (CST9431). Cells were lysed in 50 mM Tris (pH 7.5), 150 mM NaCl, 1 mM EDTA, 1% NP40 with protease and phosphatase inhibitor cocktail (Thermo Scientific), 1 mM sodium orthovanadate. Protein samples (10 mg per lane) were separated by SDS-PAGE electrophoresis and immunoblotted on polyvinylidene fluoride (PVDF) membranes and detected by ECL using a charge-coupled device (CCD) camera. Quantification was performed using ImageJ software.

Added in the revised version: Mitosis analysis

Added in the revised version: U2OS cells were synchronized at the G1 phase by treating them with 2 mM thymidine for 20 hours. Time-lapse imaging was conducted after 8∼9 hours following G1 release using an inverted confocal microscope (Stellaris; Leica) with a 20× objective (NA=0.75) using DIC mode. The images were taken in solution with different K+ concentration without serum (Fig.1A) every 10 minute. The number of mitotic cells was counted for each frame. The FBS condition included 10% FBS.

Added in the revised version: Statistical Analysis

Added in the revised version: Statistical analysis was performed using R studio. The level of statistical significance is indicated by asterisks: ns p > 0.05, * p < 0.05, ** p < 0.01, *** p < 0.001, **** p < 0.0001.